```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

# Homework 1: Introduction to Data Science (50 points)

Please complete this notebook by filling in the cells provided. Be sure to run the cell above that loads the autograder.

Note that you are not expected to understand all the Python code in this homework; you will mostly be executing code and analyzing the output, but you will not be writing extensive amounts of code.  (That will start with Homework 2 :)

Also, note that this homework touches on many of the topics we will be covering in more depth during the course, so you are not expected to fully understand the concepts yet, but answer every question to the best of your ability.  Also, several of the questions do not have a single correct answer, so apply critical thinking when answering the questions.  We will discuss the homework problems and answers in class later.

**Deadline:**

This assignment is **due Wednesday, 9/4 at 11:59pm ET**. Submissions after this time will be accepted for 24 hours and will incur a 20% penalty. Any submissions later than Thursday, 9/5 at 11:59pm ET will not be accepted unless an extension has been granted.

**Note: The tests included in this homework test only for format.  That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the [syllabus]() to learn more.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday-Friday. The schedule is posted in Canvas.

## 1. Data Visualization (10 points)
Data visualization is one of the most important subjects you will learn about in **DATA110**.  It is a great way to gain some intuition about the nature of the data; also, people are often good at discerning visual patterns.  Some questions related to data visualization are: Who is the visualization for?  What types of visualization are best for the given data?  Is there a subset of the data that might be particularly interesting to visualize?

Different types of visualization are useful for different purposes.  Some data visualization types we will be learning include scatter plots, bar plots, and histograms.

**Problem 1.1 (2 points).**  Run the code below and observe the plot that is generated. 

```python
# Just run this cell!
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
# load data file
state_data = pd.read_csv('us-state-populations.csv')
plt.xticks(rotation=-90, fontsize = 8)
sns.barplot(x=state_data['state'],y=state_data['pop_2014'])
```

Which of the following visualization types does the plot above most resemble?\
 ![](visualizations.jpg)

_Type the letter of the answer here, replacing this text_

**Problem 1.2 (2 points).**  The code in Problem 1.1 should have generated a plot with 51 vertical bars; this is called a bar plot or bar graph.  Using the generated bar plot, what 2 states would you say were closest in population to North Carolina in 2014?

_Type your answer here, replacing this text_

**Problem 1.3 (2 points).**  Run the code below and observe the plot.

```python
plt.xticks(rotation=-90, fontsize = 8)
ax=sns.barplot(x=state_data['state'],y=100*state_data['pop_2014']/state_data['pop_2014'].sum())
ax.set_ylabel('Percent of U.S. population')
```

This bar plot should look very similar to the previous; in fact, the "bar" shapes should look identical.  So how is this plot different than the previous plot?

_Type your answer here, replacing this text._

**Problem 1.4 (2 points).**  Run the code below and observe the plot.

```python
sns.histplot(state_data['pop_2014'], bins=20)
```

The plot above uses the same data, but looks very different.  This type of plot is called a *histogram*.  Note that the population label is now along the x-axis rather than the y-axis.  What do you think the "Count" label on the y-axis is referring to?  

_Type your answer here, replacing this text._

**Problem 1.5 (2 points).**  If you were asked how many states have a population under 10 million, which type of plot (bar plot or histogram) would be more useful and why?

_Type your answer here, replacing this text._

## 2. Data Representation and Types (10 points)
Statistics is the branch of mathematics related to collecting, analyzing, and interpreting data.  Data science depends heavily upon statistics.  In future homeworks you will be writing python code to do statistical analysis, but for this homework you will just be using the results from provided code to answer questions.


**Problem 2.1 (3 points)**. Consider the `state_data` table from Problem 1: 

```python
#run this code
state_data
```

Each **data sample** in this table corresponds to what?   (Use the definition of **data sample** from the slides in class)
How many data samples are there?  How many pieces of information are collected for each **data sample**?

_Type your answer here, replacing this text._

**Problem 2.2 (1 point).** Data can be re-arranged and looked at in different ways.  If all the data in the `state_data` Table was instead converted into a time-series Table, with `Year` being the first column and each data sample corresponding to a particular year, how many columns would you need for this new Table?

_Type your answer here, replacing this text._

**Problem 2.3 (4 points).** Different types of data representations are appropriate for different situations.  Which of the following data representations would you choose for each situtation, and why? (use each modality exactly once)\
**A)** Time-series table **B)** Image  **C)** Text  **D)** JSON

**Situation 1.** For every user, Netflix records all the different types of events the user might trigger, along with meta-data for the event, and new types of events and meta-data formats are created often. 

_Type your answer here, replacing this text._

**Situation 2.** For every show, Netflix has a description of the show along with the names of the actors and episode titles.

_Type your answer here, replacing this text._

**Situation 3.** For every show on Netflix there is a thumbnail to help users visualize the content.

_Type your answer here, replacing this text._

**Situation 4.** For a given user, every day Neflix logs the number of times they login, number of videos they open, and the total time spent in the application. 

_Type your answer here, replacing this text._

**Problem 2.4 (1 point).** Understanding the purposes and advantages of data types is an important aspect of being a proficient programmer.  Four basic types are `int`, `float`, `bool`, and `str` (integer, floating point, Boolean, and string).  As mentioned in class, all numeric data can also be represented as a string. So why not just make all data a string and just have one type? 

_Type your answer here, replacing this text._


**Problem 2.5 (1 point).** Similarly, Boolean values (`True`, `False`) could instead just be represented with integers with values (`1`, `0`).  What benefits can you think of for having bool as its own distinct datatype?

_Type your answer here, replacing this text._


## 3. Statistics (10 points)
Statistics is the branch of mathematics related to collecting, analyzing, and interpreting data.  Data science depends heavily upon statistics.  In future homeworks you will be writing python code to do statistical analysis, but for this homework you will just be using the results from provided code to answer questions.

**Problem 3.1 (2 points).**  Summary statistics are concise numerical measures used to help make sense out of complex data sets.  The most basic of these measures are **central tendency** measures: mean, median, and mode.  The **mean** is the arithmetic average of the data.  The **median** is the middle point of the data when arranged in ascending order.  The **mode** is the most frequently occurring value in the data.  Run the code below to compute the mean, median, and mode for the state populations in 2014, then update the code to do the same for 2017.   

```python
state_data = pd.read_csv('us-state-populations.csv')
def central_tendency_measures(table, column_name):   
       return [int(table[column_name].mean()), int(table[column_name].median()), int(table[column_name].mode()[0])]

central_tendency_measures_2014 = central_tendency_measures(state_data, 'pop_2014')
central_tendency_measures_2014
```

```python
#change the code below to compute the central_tendency_measures for 2017 using the function above
central_tendency_measures_2017 = ...
```

```python
grader.check("q1")
```

**Problem 3.2 (2 points).**
The **mode** is defined as the most common value of the data, but since the population of every state was different in 2014, the mode result isn't very useful (it just returns all the values in ascending order).  Can you think of anything you could do to the state data to make the mode function more useful for this data?  Perhaps looking at the figure in problem 1.4 will help you answer, or looking at the code below:

```python
state_data = pd.read_csv('us-state-populations.csv')
state_data['pop_2014']=(state_data['pop_2014']/1000000)
state_data['pop_2014']=state_data['pop_2014'].round()
state_data['pop_2014']=state_data['pop_2014']*1000000
new_mode = state_data['pop_2014'].mode()[0]
new_mode
```

_Type your answer here, replacing this text._

**Problem 3.3 (2 points).** Another set of statistical measures are **dispersion** measures, including range, variance, and standard deviation.  The **range** is the simplest of these; it is the difference between the maximum and minimum values.  For state population, can the mean or the median ever be bigger than the range?  Why or why not?

_Type your answer here, replacing this text._

**Problem 3.4 (2 points).**  In general, can the mean or the median ever be bigger than the range?  If so, give an example of a small (<=5 values) dataset.

_Type your answer here, replacing this text._

**Problem 3.5 (2 points).**  The **variance** and **standard deviation** are measures of how widely the data is spread around the mean.  For example, if all the data was the same, the variance would be zero, and if the data are all very different, the variance is high.  Do you think the variance of the data would be go up or down if you removed California from the dataset?  Why?

_Type your answer here, replacing this text._

## 4. Inference and Prediction (10 points)
Inference is the process of drawing conclusions from data. Prediction is the process of using the data you have to predict for cases that are not represented in the data. Inference reasons backward; from the data we have, what can we infer about what generated that data?  Prediction reasons forward; given the data we have, what can we predict about other data?  They are related concepts, but distinctly different.

**Problem 4.1 (2 points).** Consider a situation in which you receive a quarter from a vending machine.  If you toss the quarter 20 times, what would be your prediction for how many heads and how many tails you would get?  What would you guess the probability would be for getting exactly this many heads and tails? (don't do math for this, just give an estimate)


_Type your answer here, replacing this text._

**Problem 4.2 (2 points).** Suppose you actually do toss the quarter 20 times, and you get 15 heads and 5 tails.  Can you infer anything about the properties of the quarter?  Why or why not?


_Type your answer here, replacing this text._

**Problem 4.3 (2 points).** Suppose you toss the quarter 2000 times, you get 1500 heads and 500 tails.  Is this situation any different than **Problem 4.2**?  If so, what is different?

_Type your answer here, replacing this text._

For coin tosses or dice rolls, most people have a good mental model and can interpret the data without sophisticated tools.  But for more complex or less familiar data, inferences and predictions are much less obvious.  That is when data analysis tools such as **linear regression** are particularly useful.  We'll learn more about linear regression later; for now, just think of it as fitting a model that describes the data.

**Problem 4.4 (2 points).** A random sample of 10 used Teslas that are for sale results in the data shown in the code below.  Run the code to see a table of the sales data. 

```python
tesla_prices = pd.read_csv('tesla_list.csv')
tesla_prices
```

The table data can be challenging to understand by just looking, so plotting the data is always a good idea.  Run the code cell below to see a scatter plot of Price vs. Age.

```python
sns.scatterplot(data=tesla_prices, x="Age", y="Price")
```

The scatterplot is helpful, but by performing a linear regression and visualizing the model that is fit, we can get even more insight into the data.  Run the code cell below then answer the questions.

```python
g=sns.regplot(data=tesla_prices, x="Age", y="Price")
plt.grid()
plt.show()
```

Based on the model graphics you see in the plot above, would you **infer** that price variablility in dollars is greater for **A. newer** or **B. older** Teslas?  

_Type your answer here, replacing this text._

**Problem 4.5 (2 points).** Based on the model graphics in the plot above, what would you **predict** for the price of a 8 year Tesla?


_Type your answer here, replacing this text._

## 5. Ethics (10 points)

Continually throughout the semester, we will explore the ethical considerations for data science.  This is because ethical concerns can arise at every stage of data science; for example, during data collection, data processing, data storage, data analysis, data visualization, and even data destruction. 

**Problem 5.1 (4 points).**
Suppose you are tasked with doing a study of the effects of diet on sleep, so you log your daily calories and sleep hours for a month.  At the end of the month, you realize you forgot to record sleep data for a few of the days and calorie data for a few other days.  Order the four behaviors below from most appropriate to least appropriate and explain any ethical concerns you have with each behavior.
1. Throw out all the data and start the study again.  
2. Do your best to remember what the values were for the data you forgot to record, and add those values into the collected data.
3. Use the days that have both data recorded to predict the missing values in the collected data and add those values to the collected data.
4. Use only the days that have both data recorded for analysis.

_Type your answer here, replacing this text._

**Problem 5.2 (2 points).**  Suppose you decide to do an extended study of the effects of diet on sleep, so you collect and analyze six months of data, and write and submit a paper about your findings.  While writing the paper, you gather 1 more month of data, but this data doesn't show the same relationships as the first six months.  Do you decide to include it in the paper?  Why or why not?

_Type your answer here, replacing this text._

**Problem 5.3 (2 points).** While making a visualization for your paper about the effects of diet on sleep, you plot the data and see that one data point for calories is 10 times more than any other day.  You are fairly sure you did not eat 10 times more calories that day.  Do you handle this data point differently?  Why or why not?

_Type your answer here, replacing this text._

**Problem 5.4 (2 points).** Just as you finish writing your paper about the effects of diet on sleep, you accidentally save over the file that contains the raw data you collected, and unfortunately you do not have a backup copy.  The paper is finished, but you can no longer reproduce the results it describes.  Do you still submit it for publication?  If so, do you change the paper in any way?

_Type your answer here, replacing this text._

## Bonus.  Images (? extra points)
Though we won't do a lot with images in this class, Python does provide tools for working with images, and images are very similar to Tables.  Run the code below that loads an image and draws a circle in the image.

```python
img = plt.imread('yikes.jpg')
img = img/255

def draw_circle( x0: int, y0: int, q: int, t: int=10, c: float=1.0):
    for i in range(0,360):
        x = x0 + np.cos(np.radians(i)) * q
        y = y0 + np.sin(np.radians(i)) * q
        x=int(x)
        y=int(y)
        img[y - t: y + t, int(x) - t : x + t ]=c
draw_circle( 100, 100, 50, 3, 0 )
plt.imshow(img)
```

Who said the instructor's cat was cute?!?  

Now write code to call the `draw_circle` function so that a white circle is drawn around the cat's head. You don't need to understand the Python code or the math behind it to figure out what each of the parameter does! Play around with the function by passing in different numbers and see how that changes what you see. For example, what does `draw_circle(1000, 100, 50, 3, 0)` do? What did that change? Remember, you can add as many additional cells as you'd like.

Copy the following code into the next cell and change the line marked with `...`.

```Python
img = plt.imread('yikes.jpg')
img = img/255
# Change the code below to draw a white circle around the cat's head
draw_circle(...) 
plt.imshow(img)
```

```python
# copy and edit code here and run cell

```

Finally, load an image of your own and draw circles on it, or other shape functions you write!  Hint: we like pets. 

Copy the following code into the next cell and change the line marked with `...`.

```Python
# Change the code below
# Make sure to write the correct file name here for your image.
img = plt.imread(...)
img = img/255
# Change the code below to draw a white circle around the cat's head
draw_circle(...) 
draw_circle(...) 
plt.imshow(img)
```

```python
# copy and edit code here and run cell

```

You're all done with Homework 1!

**Important submission steps:**

1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**.
3. Click the link to download the zip file.
4. Go to Canvas and submit the zip file.
**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Gradescope Homework1

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert central_tendency_measures_2017 == [6386650, 4454189, 579315]
```
