```python
# Initialize Otter
import otter
grader = otter.Notebook("hw2.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)

Names and Onyens of fellow students you discussed Homework2 problems and ideas:

- Students: 

We encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 2: Python and Pandas (100 points)

Please complete this notebook by filling in the cells provided. Be sure to run the cell above that loads the autograder, and be sure to save your notebook before generating your submission by running the last cell.

!! Please submit a PDF of your notebook and upload it to Gradescope along with the generated zip file.  If no PDF submitted, **points will be deducted**.  You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop.  You can then drag and drop the zip & pdf file together to the upload box on Gradescope.

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!

Note that for homework, the autograder is useful feedback that you should pay attention to, but passing the autograder for a problem **DOES NOT GUARANTEE** you will get full credit on the problem, and some problems will be manually graded.


## Problem 1. Python Lightning Round (15 points)
For each of the sentences below, write the corresponding Python code.  For example, for the sentence "Assign the string `Rick is left-handed` to the variable named `inconceivable`", the corresponding code would be:

`inconceivable = 'Rick is left-handed'`

Don't change any of the code that is provided.  Note that answers may require more than a single line of code.

**Problem 1.1 (1 point).** Assign 4.5 to a variable named `start_off_easy`.

```python
import numpy as np
import pandas as pd
```

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.1")
```

**Problem 1.2 (1 point).** Assign 5 to the variables named `x` and `y` then assign the sum of `x` and `y` to a variable named `z`.

```python
...
```

```python
grader.check("q1.2")
```

**Problem 1.3 (1 point).** To a variable named `my_stuff`, assign a list that contains 5, the string 'pencil', and -4.2 (in this order).

```python
...
```

```python
grader.check("q1.3")
```

**Problem 1.4 (1 point).** Assign a range of integers to a variable named `bizz` that starts at 3 and steps by 3 up to and including 12.

```python
...
```

```python
grader.check("q1.4")
```

**Problem 1.5 (1 point).** Define a function named `test_func_prob15` that takes in a list `exactly_three_things`, and assigns its **third** element to a variable named `favorite`. The function then returns `favorite`.

```python
def test_func_prob15(exactly_three_things):
...
    return favorite
```

```python
grader.check("q1.5")
```

**Problem 1.6 (2 points).** Define a function named `take_away` that takes in two numbers `x`, `y`, and returns the first argument minus the second argument.

```python
...
```

```python
grader.check("q1.6")
```

**Problem 1.7 (2 points).** Inside the function, if the given variable `numeric_arg` is greater than 0, set a variable `positive_flag` to `True`.

```python
def test_positive(numeric_arg):
    positive_flag = False
...
    return positive_flag
```

```python
grader.check("q1.7")
```

**Problem 1.8 (2 points).** Loop through all the elements in a list of strings named `lots_of_words` and increment a variable name `word_count` by a value of 1 for each element. 

**Hint**: Check your code by comparing the returned variable `word_count` with the total number of workds in your list.

```python
def test_func_prob18(lots_of_words):
    word_count = 0
...
    return word_count
  
```

```python
grader.check("q1.8")
```

**Problem 1.9 (2 points).** Return the second to the last item in a list named `many_things`.

```python
def test_func_prob19(many_things):
...
```

```python
grader.check("q1.9")
```

**Problem 1.10 (2 points).** Define a function `fun_game` with one argument `i` that returns the string 'bizz' if the argument is divisible by 3.  If not, it returns the string 'buzz' if the argument is divisible by 7.  If the argument is not divisible by either 3 or 7, the function returns string 'boring'.

```python
...
```

```python
grader.check("q1.10")
```

**Extra credit (1 point).** Wait, already?!?  Define a function `funner_game` with one argument that returns the string 'bizz' if the argument is divisible by 3, the string 'buzz' if the argument is divisible by 7, unless it is divisible by both 3 and 7, in which case it returns string 'bizzbuzz'.  If the argument is not divisible by either 3 or 7, the function returns string 'boring'.

```python
...
```

```python
grader.check("q1.11")
```

## Problem 2. Rideshare (30 points)

### Part A: Reviewing functions and for loops
 
**Problem 2.1 Squared lists (1 point).** 
Write a function called `square` which takes in a list of numbers and returns a list with each item squared. Lecture 5 slides might be helpful if you're stuck.

```python
def square(L):
    ...
```

```python
grader.check("q2.1")
```

**Problem 2.2 Power lists (2 points).** 
Now, let's make a more flexible function called `compute_power(L,n)` that lets the user raise each element to the `n`th power. Give `n` a default value of 2. Test by checking that `square(L) = compute_power(L)` for different examples of `L`.

```python
def compute_power(L,n=2):
    ...
```

```python
grader.check("q2.2")
```

### Part B: Computing distances with functions

For the next part of this assignment, we're going to think about how we might be able to write functions that can be used to compute distances on a map.

**Problem 2.3 Computing the Euclidean distance (3 points).** 

"Euclidean distance" is simply the distance of the straight line that connects two points in Euclidean space. It's basically the distance "as the crow flies". On a map, it might look something like this:

<img src="https://i.imgur.com/KSmguA5.png" width=500px>
<div style="text-align: center"> </div>

Write a function called `Euclidean_distance` which takes in two points `p` and `q` (both are lists of numbers) and returns the calculated Euclidean distance bewteen `p` and `q`.

For example: 

```python
Euclidean_distance([2, 1], [5, 5])
```
```
5.0
```


**Hint:** Euclidean distance between $(x_1,y_1)$ and $(x_2,y_2)$ is defined as $\sqrt{(x_1-x_2)^2+(y_1-y_2)^2}$. (Remember the Pythagorean theorem?)

**Hint:** The Python code for calculating square root of `x` is the following:

```python
import math
math.sqrt(x)
```
We have already imported math for you.

```python
import math

...
```

```python
grader.check("q2.3")
```

**Problem 2.4 Computing the "Lyft" distance (3 points).** 

For this part of the assignment, let's say that you've been thinking about starting a new rideshare company similar to Uber or Lyft. You've started thinking about how you should go about estimating the amount you should charge for a given ride. You've realized that computing the Euclidean distance between two points makes sense for "as the bird flies" distance computations, but not for vehicles which need to navigate city blocks. This is why this type of distnace is called the Manhattan distance. Now you need to create a new function that computes an estimate for the total distance of the ride for your new "Lyft-like" service.

When you're driving a car you can't just travel in a straight line between two points, you're forced to drive along the grid structure defined by where the streets are. Your path might look something like this:

<img src="https://i.imgur.com/ieJ1Z4k.png" width=500px>
<div style="text-align: center"> </div>
<br>

For the purposes of computing these new "Lyft"-like distances, you're going to assume that the drivers for your company are generally pretty good about taking the shortest possible route, and you're also going to assume that, most of the time, the streets will have a pretty standard grid-like structure. After all, it would be pretty hard to predict just how irregular the routes might end up being.

**Write a function for computing a "Lyft" distance** based on the restriction that the car must travel on a grid of roads that run parallel to the $x$ and $y$ axes of an $x$-$y$ plane. Your function should be named `lyft_distance` and take a **starting point and a stopping point as arguments** and **return** the distance estimate. Your estimate will essentially tell you how many blocks the driver will travel to get to the destination. 

**Hint**: "Lyft" distance between $(x_1,y_1)$ and $(x_2,y_2)$ is defined as $\vert x_1-x_2\vert + \vert y_1-y_2\vert$.

**Hint**: The Python syntax for absolute value is `abs(x)`.


```python
...
```

```python
grader.check("q2.4")
```

**Problem 2.5 Creating a fare-computing function (4 points).** 

Let's assume that you want to be able to quickly calculate an estimate for the cost of the trip so that you can let your riders know roughly how much you'll charge them.  

**Write another function** called `trip_cost` that **calls the function you just wrote** for computing the total distance of the ride to ***return* the cost of the trip in a user-friendly sentence** so that it reads "The cost of the trip will be $< amount >." where < amount > is the total in dollars that the trip will cost. Remember that we only want numbers up to 1 cent (no fractional cents). Your function should take as arguments the starting point, the stopping point, and the fare (per mile), with a default value of 40 cents. (For the purposes of this exercise, we will assume that each block is one mile so that we don't need to worry about unit conversions.)

***Hints:***

1.) Start with the `def` statement. How many parameters do you need? Do any of them have a default value? Next, think about what you need to `return`. What should the data type be?

2.) To make the dollar value look nice, first calculate the cost, and then generate a string using the code `f"The cost of the trip will be ${cost:.2f}."` Wow, what is this pretty fancy string format? They are called Python f-string, and you can read about them here (https://realpython.com/python-f-strings/) and here (https://docs.python.org/3/tutorial/inputoutput.html).

3.) Your `trip_cost` function should call your `lyft_distance` function.

```python
...
```

```python
grader.check("q2.5")
```

### Part C: Combining your functions

Now that you've written two different functions for finding the distance, you're thinking that it would be more convenient if you could use the same function to compute both the Euclidean distance and the Lyft distance. It would also make it a lot easier to compare how much longer the Lyft distance is than the Euclidean distance.

**Problem 2.6 (5 points).**

**Write a function** called `calculate_distance` that is capable of computing the Euclidean distance **or** the Lyft distance based on an argument called `path_type`. 
* If the `path_type` is 'Lyft', compute the Lyft distance, otherwise compute the standard Euclidean distance. 'Lyft' should be your default parameter value.
* Your function should call the `lyft_distance` and `Euclidean_distance` functions that you have written earlier.

As before, your function should **return** the distance.


```python
...
```

```python
grader.check("q2.6")
```

### Part D: Planning your route

Now that you've worked out a system for estimating travel distances and ride fares, you're ready to start testing your new rideshare service. You fire up your fancy new program and immediately get three simultaneous rider requests, but luckily they all are looking for a ride to the same place -- the airport!

Let's assume that:

1. Your starting location is (-2,3).
2. The ride locations are (-1, -2), (3,3), and (2, -1).
3. The stopping location, the airport, is at (4,0).

Run the cell below to assign the locations that are used for the rest of this section.  Note that _stops_ is actually a list of lists!  

```python
start = [-2, 3]
stop0 = [-1, -2]
stop1 = [3, 3]
stop2 = [2, -1]
stops = [stop0, stop1, stop2]
end=[4,0]
stop0 = [] # clear these, we will only used stops in our code
stop1 = [] # reminder about list assignment:
stop2 = [] # since these are cleared after stops is assigned, clearing them doesn't affect stops
stops
```

Run the cell below to visualize the locations.  Note: you are not expected to understand the code in `plot_stops()`

```python
from matplotlib import pyplot as plt

def plot_stops(start, stops, end):
    x_stops=[]
    y_stops=[]   
    for i, stop in enumerate(stops):
        x_stops.append(stop[0])
        y_stops.append(stop[1])
        plt.text(stop[0]-0.7,stop[1], f'Stop {i}')

    plt.scatter(x_stops,y_stops)
  
    plt.scatter(start[0],start[1],color="red")
    plt.text(start[0]+0.5,start[1]-0.1, 'Start')
    plt.scatter(end[0],end[1],color="gold")
    plt.text(end[0]-0.7,end[1], 'End')

plot_stops(start,stops,end)
```

Let's first assume that we pick the riders up in order, i.e. that we go first from the start to (-1, -2), then from (-1,-2) to (3,3), then from (3,3) to (2, -1), and then lastly from (2,-1) to the airport.

**Problem 2.7 (5 points).**

Write a function called `simple_route3` which inputs three variables, `start`, `stops`, and `end`, and then returns the total distance of the trip. Your solution should call the `lyft_distance` function several times. You can assume that there will always be 3 stops. 

```python
...
```

```python
grader.check("q2.7")
```

Now, lets figure out the best order to pick the riders up.  To do that, we are going to call simple_route3 with all the possible orderings of the riders.  For example, we already examined the ordering `[stop0, stop1, stop2]`.  Another ordering would be `[stop0, stop2, stop1]`.  These orderings are also known as _permutations_.  For three riders, there are 6 unique orderings or _permutations_.

**Problem 2.8 (1 point)** There are shortcuts in Python for generating permutations, but we'll code them by hand for now.  Finish appending the the six different permutations of the stops to list variable named _perms_.  (Note that _perms_ is a list of lists of lists!  That's a lot of square brackets!)

```python
def list_of_perms3(stops):
    perms3 = []
    perms3.append([ stops[0], stops[1], stops[2] ])
    perms3.append([ stops[0], stops[2], stops[1] ])
...
    return perms3
```

```python
grader.check("q2.8")
```

**Problem 2.9 (5 points).**

**Now for the most challenging part**: You need to create a function `best_route3` that takes in a starting position, a list of **three** rider locations, and a stopping location and determines the order of rider pickups that will require the least amount of driving. 

Once you've come up with a plan, try **writing the function**. The function should **return the list of rider locations in the appropriate pick-up order**.

Your function should have three input parameters start, stops, and end. It should return a list of the 3 stops in the best order to minimize driving.   It is recommended that your `best_route3` function call your `simple_route3` function and `list_of_perms3` function that you wrote earlier, although this is not required. 


**Hint**: There are many different ways to do this, most of which will involve a `for` loop.  One way would be to create a list of all the distances, and then use `min()` function and the `list.index()` function to get the index of which distance is minimum. You can assign as many variables and write as many lines as you need inside the function.

```python
...
```

```python
best_perm = best_route3(start,stops,end)

plot_stops(start, stops, end)

def plot_line(p,q):
    plt.plot([p[0],q[0]],[p[1],q[1]],color="black")
    
plot_line(start,best_perm[0])
plot_line(best_perm[0],best_perm[1])
plot_line(best_perm[1],best_perm[2])
plot_line(best_perm[2],end)

```

```python
grader.check("q2.9")
```

Now run the code below to visualize your best route. (It is okay if you don't understand the code yet.)

**Problem 2.10 (1 points).** Using the functions you've defined, how would you calculate the distance of the best route given start, stops, and end?  

**Hint**: for the points from the start of Part D, the answer should be 19!

```python
def best_distance3( start, stops, end):
...

best_distance3( start,stops,end)

```

```python
grader.check("q2.10")
```

## Problem 3. Creating Tables (15 points)
In most cases, you will be using data that you are reading in from a file; in fact, sometimes you might be using data that is so big it can't fit in memory, and HAS to be accessed from a file.  But it can still be useful to know how to build a table from scratch, and understanding how to do it also provides insight into the behavior of tables.  Run the code below to see an example of how to create a table. 

```python
import pandas as pd
row_list0 = [1, 2, 3]
row_list1 = [4, 5, 6]
row_list2 = [7, 8, 9]
df = pd.DataFrame( data=[row_list0, row_list1, row_list2], columns=['name', 'rank', 'serial number'])
df
```

**Problem 3.1 (3 points).** Suppose you're using Python to help solve a high school Algebra problem.  You have 4 apples, 3 oranges, and 3 pineapples.  Create a table (pandas DataFrame) named `fruits` that contains this information. It should have two attributes or pieces of information about each fruit: `'Fruit name'` and `'Amount'`.  Note: Use lower-case and singular words for the name of each fruit, like "apple".  

```python
...
answer31=fruits.copy() # answer31 is used by the auto-grader
```

```python
grader.check("q3.1")
```

**Problem 3.2 (2 points).** Write a single line of code that assigns the total of the number of pieces of fruit in the table to a variable named `pieces_of_fruit`.

```python
pieces_of_fruit = ...
pieces_of_fruit
```

```python
grader.check("q3.2")
```

**Problem 3.3 (2 points).** Add a new attribute to `fruits` named `'Price'` and initialize the price for all the fruits to be 0.99.

```python
...
answer33 = fruits.copy() # answer33 is used by the auto-grader. do not change this line!
```

```python
grader.check("q3.3")
```

**Problem 3.4 (2 points).** Now, let's add another row to the table.  To do this, note that `df.loc[rownum]` returns a particular row of `df` according to the row label (index), and `len(df)` returns the current number of rows of `df`.  To add a new row to our example table from above:    

```python
df.loc[len(df)] = [ 10, 11, 12]
df
```

Your turn.  Suppose you have a second batch of 6 oranges priced at 0.99, add those to the end of the table `fruits`.

```python
...
answer34 = fruits.copy() # answer34 is used by the auto-grader. do not change this line!
```

```python
grader.check("q3.4")
```

**Problem 3.5 (2 points).** Now let's change the price for just oranges.  First, lets get the rows that corresponds to oranges using **Boolean indexing** as covered in lecture.  Assign them to the variable `orange_rows`.

```python
orange_rows = ...
orange_rows
```

```python
grader.check("q3.5")
```

**Problem 3.6 (2 points).** You could also use **Boolean indexing** as covered in lecture to get all the rows that are not oranges.  Assign those rows to the variable `not_orange_rows`.

```python
not_orange_rows = ...
not_orange_rows
```

```python
grader.check("q3.6")
```

**Problem 3.7 (2 points).** Okay, back to changing the price for oranges.  To do this, we need to set the value of particular cells.  We used `loc()` to add a new row; similarly, we can use `loc()` to assign cells from a particular column for our Boolean indexed rows.  Again, see the example for `df`: 

```python
df.loc[ df['name'] == 4, 'serial number' ] = 99
df
```

The above code is effectively "For all the rows in which `name` equals 4, set the value for `serial number` to be 99.  Similarly, now you set the price in `fruits` for just oranges to 2.99 using Boolean indexing and `loc()`.

```python
...
fruits
```

```python
grader.check("q3.7")
```

## Problem 4. R.O.U.S.es (20 points)
**R**odents **O**f **U**nusual **S**ize are fabled to be one of the most fearsome creatures you might ever encounter, but accounts of their attributes vary widely.  Run the following code to load a table of all known data about R.O.U.S.es.

```python
import pandas as pd
import seaborn as sns
rodents_of_unusual_size = pd.read_csv("ROUSes.csv")
rodents_of_unusual_size

```

**Problem 4.1 (2 points).** How many attributes or different pieces of information have been measured for R.O.U.S.es?  Assign that number to `num_attributes`.

```python
num_attributes = ...

```

```python
grader.check("q4.1")
```

**Problem 4.2 (2 points).** One of the best ways to get a feel for data is to visualize it.  Using seaborn's `scatterplot()` function that you learned in class, draw a plot with `Age` on the x-axis, and `Length` on the y-axis.

```python
plt = ...
```

**Problem 4.3 (2 points).** Draw another scatterplot of `Age` vs. `Weight`.

```python
plt = ...
```

**Problem 4.4 (2 points).** Finally, do one more scatterplot of `Length` vs. `Weight`.

```python
plt = ...
```

**Problem 4.5 (2 points).**  What general relationships, if any, do you see in the numeric data?  For example, do the data seem completely unrelated to each other?  Or is one attribute increasing proportionally to another, or is inversely proportional (i.e. decreases while the other one increases)?  Without enough data, sometimes it can be hard to see higher order relationships, like one attribute depending on the square of another, but maybe you can guess at one such a relationship in this data? No calculation needed.

_Type your answer here, replacing this text._

**Problem 4.6 (2 points)** For non-numeric data, graphical visualization isn't as helpful, but sometimes sorting or grouping the data by an attribute can help.  Sort the data by `Age` and look at the data table again. No need to save or assign the modified table to a variable.

```python
...
```

**Problem 4.7 (2 points).**  With the data sorted by age, what trends do you see in the `Temperament`?  Are there any exceptions to the trends? No need to calculate anything.

_Type your answer here, replacing this text._

**Problem 4.8 (3 points).** Some people report that real danger from R.O.U.S.es is that they are very uncoordinated and trip and fall onto people.  If `Weight` poses the biggest threat, on average what `Temperament` would be the most dangerous to run into?  Write a function `average_weight_by_temperament` that calculates the mean of the weight for each `Temperament`, and **print**s it along side the `Temperament`.  This function does not return anything.

Hint: this is just like a problem we covered in lecture.

This question is manually graded.

```python
def average_weight_by_temperament( df ):
...
average_weight_by_temperament( rodents_of_unusual_size )
```

**Problem 4.9 (2 points).**  Sometimes even a single data point can really change the statistics of a dataset, so it can be useful to examine such individual data points.  Which data point is most extreme in this dataset?  Is this data point still consistent with the general trends in the data?  Using mental math, if we removed this point, and recalculated the means, what effect would it have on the ordering of dangerousness?

_Type your answer here, replacing this text._

**Problem 4.10 (1 points).**  What key descriptor is missing from the data in the table that might help you decide if R.O.U.S.es (assuming they exist) pose a massive threat?  

**Hint**: the answer we're looking for is not an additional column/new attribute.

_Type your answer here, replacing this text._

## Problem 5. (20 points). Burritos
A certain Carolina student from DATA110 is visiting California for Fall Break this year, and they intend to use their newly-mastered Data Science skills to make sure they have some great burritos while they are there.  (Whereas the other DATA110 students just get this delightful homework problem out of the experience.)  The student found some data files with reviews and menu information for some California burrito restauraunts, but unfortunately the data was not nearly as clean as most of what they had dealt with in class, so they had to do some work on the data first.

**Problem 5.1 (2 points).** Write the code to load the file "ratings.csv" into a DataFrame named `ratings`. In the next line, write the code to load the file "burritos_types.csv" into a DataFrame named `burritos_types`.

```python
import pandas as pd
pd.options.display.max_colwidth=15

ratings = ...
burritos_types = ...

print(ratings)
print()
print(burritos_types)
```

```python
grader.check("q5.1")
```

**Problem 5.2 (1 points).** It would be easier to deal with just one table, so next you will combine the two tables above into a single table named `burritos`.  To do this, it helps to find an attribute that is common between the 2 tables.  What attribute is that for these tables?  

```python
common_attribute = ...

```

```python
grader.check("q5.2")
```

**Problem 5.3 (2 points).**  The student is only interested in restaurants that have both a rating and some menu information, so they want to combine the tables using what is called an "inner join".  Basically, for an inner join, the common attribute needs to appear in both tables.  There are other kinds of join, but inner join is one of the most useful.  The join() function syntax is a litte tricky, so here is code to do the join.

```python
burritos = burritos_types.join(ratings.set_index('Name'), on='Name', how='inner')
burritos2 = ratings.join(burritos_types.set_index('Name'), on='Name', how='inner')
print(burritos)
print()
print(burritos2)
```

When joining 2 tables together, one table is actually "joining" the other, and the results are slightly different depending on which one you have join the other.  As you can see from the output above, the shape of the tables is the same, but the order of the attributes is different.  The attributes of the table that is joining come after the table it is joining (except for the common attribute, which is already in the first table).  By looking at the output above, What other difference do you notice in addition to attribute order?

_Type your answer here, replacing this text._

**Problem 5.4 (2 points).** The student decides to use `burritos` for the rest of their analysis.  Write code to calculate average cost of items in the table `burritos` rounded to the nearest cent?

```python
average_cost = ...

```

```python
grader.check("q5.4")
```

**Problem 5.5 (2 points).** Make a new table called `top_rated_burritos` that is just like `burritos` but only includes entries that have a Google rating >= 4.7.

```python
top_rated_burritos = ...
print(top_rated_burritos)


```

```python
grader.check("q5.5")
```

**Problem 5.6 (2 points).** Write pandas code to compute the average cost of items entries that have a Google rating >= 4.7, again rounded to the nearest cent.

```python
average_cost_top_rated = ...
print(average_cost_top_rated)


```

```python
grader.check("q5.6")
```

```python
print(round(average_cost_top_rated-average_cost,2))
```

**Problem 5.7 (2 points).** Only 62 cents more for a great burrito instead of an average one...count me in!  But then the student realized, not all the items are the kind of burrito that they like, or even burritos at all!  In particular, the student likes California burritos, which just have a menu name containing the word `California`.  To make sure there are enough choices, the student also decides to lower their standards a little on the rating requirement.  Write code to calculate how many California burrito options there are that have a Google rating >= 4.5.

**Hint**: First figure out how to check if a word is in a string (there are several ways to do this). Take a look at [https://pandas.pydata.org/docs/reference/api/pandas.Series.str.find.html](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.find.html) for one such way.

**Hint**: Your answer does not have to be in a single ine of code.

```python
top_rated_Califoria_count = ...
print(top_rated_Califoria_count)
```

```python
grader.check("q5.7")
```

**Problem 5.8 (2 points).** Okay, that is a nice number of options. But now the student starts to question whether they should be trusting Google for the reviews.  They decide it would be good to see a how Yelp reviews compare to Google reviews, so they decide to make a scatter plot with Yelp scores on the x-axis and Google scores on the y-axis.   Write the code to do this for the table `burritos`.  (We've also plotted a line to help with the next question.) 

```python
plt = ...
plot_line( [2,2], [5,5])
```

**Problem 5.9 (1 points).** What do you notice about how the ratings compare?  For example, is one generally lower than the other?  Does one have more precision than the other?   

_Type your answer here, replacing this text._

**Problem 5.10 (1 points).** Now, feeling confident in their analysis, the student feels they are armed with the knowledge they will need to have a great burrito experience.  But when their DATA110 professor finds out about what they've been working on, the professor asks them to clean up the data even more so it can be used for future DATA110 classes.  First, they ask the student to make a new table `burritos_compact` that removes the confusing 'Unnamed: 0' attribute.  Write the code below.

```python
burritos_compact = ...
print( burritos_compact)
```

```python
grader.check("q5.10")
```

**Problem 5.11 (2 points).** And finally, the professor asks the student to make a new table named `burritos_extravalue` from a copy of `burritos_compact`, and adds one extra attribute.  The attribute should be named `'Value'` which is 100.0 times the Google rating for a restaurant divided by the item cost, rounded to the nearest integer.  Write the code for `burritos_extravalue` below.

```python
burritos_extravalue = burritos_compact.copy()
...
print(burritos_extravalue)
```

```python
grader.check("q5.11")
```

**Extra credit (1 points).** When you create `burritos_extravalue`, can you make the data in the `Value` attribute be of type `int` rather than type `float`?

```python
burritos_extravalue = burritos_compact.copy()
...
print(burritos_extravalue)
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework2

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1 (points: 1)
```
>>> assert start_off_easy == 4.5
```


Tests for q1.10 (points: 2)
```
>>> assert fun_game(60) == 'bizz'
>>> assert fun_game(140) == 'buzz'
>>> assert fun_game(210) == 'bizz'
>>> assert fun_game(500) == 'boring'
```


Tests for q1.11 (points: 1)
```
>>> assert funner_game(60) == 'bizz'
>>> assert funner_game(140) == 'buzz'
>>> assert funner_game(210) == 'bizzbuzz'
>>> assert funner_game(500) == 'boring'
```


Tests for q1.2 (points: 1)
```
>>> assert x == 5
>>> assert y == 5
>>> assert z == 10
```


Tests for q1.3 (points: 1)
```
>>> assert my_stuff == [5, 'pencil', -4.2]
```


Tests for q1.4 (points: 1)
```
>>> assert type(bizz) == range
>>> assert len(bizz) == 4
>>> assert bizz[0] == 3
>>> assert bizz[1] == 6
>>> assert bizz[2] == 9
>>> assert bizz[3] == 12
```


Tests for q1.5 (points: 1)
```
>>> assert test_func_prob15(['Whiskers', 'on', 'kittens']) == 'kittens'
```


Tests for q1.6 (points: 2)
```
>>> assert take_away(10, 6) == 4
>>> assert take_away(20, 16) == 4
```


Tests for q1.7 (points: 2)
```
>>> assert test_positive(10) == True
>>> assert test_positive(-3) == False
>>> assert test_positive(0) == False
```


Tests for q1.8 (points: 2)
```
>>> assert test_func_prob18(['Whiskers', 'on', 'kittens']) == 3
>>> assert test_func_prob18([]) == 0
```


Tests for q1.9 (points: 2)
```
>>> assert test_func_prob19(['Whiskers', 'on', 'kittens']) == 'on'
>>> assert test_func_prob19(['Whiskers', 'on']) == 'Whiskers'
```


Tests for q2.1 (points: 1)
```
>>> assert square([25, 29, 30]) == [625, 841, 900]
```


Tests for q2.10 (points: 1)
```
>>> assert best_distance3(start, stops, end) == 19
```


Tests for q2.2 (points: 2)
```
>>> assert compute_power([23, 28, 20]) == [529, 784, 400]
>>> assert compute_power([-2, 28, 20], 5) == [-32, 17210368, 3200000]
```


Tests for q2.3 (points: 3)
```
>>> assert Euclidean_distance([2, 1], [5, 5]) == math.sqrt((2 - 5) ** 2 + (1 - 5) ** 2)
>>> assert Euclidean_distance([22, 10], [-2, 3]) == math.sqrt((22 + 2) ** 2 + (10 - 3) ** 2)
>>> assert Euclidean_distance([-2, 3.2], [2.8, -3.9]) == math.sqrt((-2 - 2.8) ** 2 + (3.2 + 3.9) ** 2)
```


Tests for q2.4 (points: 3)
```
>>> assert lyft_distance([4, 5], [-2, -6]) == abs(4 + 2) + abs(5 + 6)
>>> assert lyft_distance([2.1, -1.98], [3, 2]) == abs(2.1 - 3) + abs(-1.98 - 2)
```


Tests for q2.5 (points: 4)
```
>>> assert trip_cost([-2, -6], [4, 5]) == 'The cost of the trip will be $6.80.'
>>> assert trip_cost([3, 98], [2, -1.3]) == 'The cost of the trip will be $40.12.'
```


Tests for q2.6 (points: 5)
```
>>> assert calculate_distance([1, 3], [6, 1]) == 7
>>> assert calculate_distance([1, 3], [6, 1], 'Euclidean') == math.sqrt(29)
```


Tests for q2.7 (points: 5)
```
>>> assert simple_route3(start, stops, end) == 23
>>> assert simple_route3([4, 3], [[3, 6], [4, 4], [3, 3]], [2, 2]) == 11
```


Tests for q2.8 (points: 1)
```
>>> assert len(stops) == 3
>>> perms = list_of_perms3(stops)
>>> assert len(perms) == 6
>>> assert [stops[1], stops[0], stops[2]] in perms
>>> assert [stops[1], stops[2], stops[0]] in perms
>>> assert [stops[2], stops[0], stops[1]] in perms
>>> assert [stops[2], stops[1], stops[0]] in perms
```


Tests for q2.9 (points: 5)
```
>>> assert best_route3(start, stops, end) == [stops[0], stops[2], stops[1]]
```


Tests for q3.1 (points: 3)
```
>>> assert answer31.equals(pd.DataFrame([['apple', 4], ['orange', 3], ['pineapple', 3]], columns=['Fruit name', 'Amount'])) == True
```


Tests for q3.2 (points: 2)
```
>>> assert pieces_of_fruit == 10
```


Tests for q3.3 (points: 2)
```
>>> assert answer33.equals(pd.DataFrame([['apple', 4, 0.99], ['orange', 3, 0.99], ['pineapple', 3, 0.99]], columns=['Fruit name', 'Amount', 'Price']))
```


Tests for q3.4 (points: 2)
```
>>> assert answer34.equals(pd.DataFrame([['apple', 4, 0.99], ['orange', 3, 0.99], ['pineapple', 3, 0.99], ['orange', 6, 0.99]], columns=['Fruit name', 'Amount', 'Price']))
```


Tests for q3.5 (points: 2)
```
>>> assert orange_rows.equals(answer34[answer34['Fruit name'] == 'orange'])
```


Tests for q3.6 (points: 2)
```
>>> assert not_orange_rows.equals(answer34[answer34['Fruit name'] != 'orange'])
```


Tests for q3.7 (points: 2)
```
>>> assert fruits.equals(pd.DataFrame([['apple', 4, 0.99], ['orange', 3, 2.99], ['pineapple', 3, 0.99], ['orange', 6, 2.99]], columns=['Fruit name', 'Amount', 'Price']))
```


Tests for q4.1 (points: 2)
```
>>> assert num_attributes == 4
>>> assert rodents_of_unusual_size.shape[0] > 0
```


Tests for q5.1 (points: 2)
```
>>> assert ratings.shape == (82, 4)
>>> assert burritos_types.shape == (249, 4)
```


Tests for q5.10 (points: 1)
```
>>> assert burritos_compact.shape == (212, 6)
```


Tests for q5.11 (points: 2)
```
>>> assert burritos_extravalue.iloc[0, 6] == 68
>>> assert burritos_extravalue.iloc[2, 6] == 65
```


Tests for q5.12 (points: 1)
```
>>> import numpy as np
>>> assert type(burritos_extravalue.iloc[2, 6]) == np.int64 or type(burritos_extravalue.iloc[2, 6]) == int
```


Tests for q5.2 (points: 1)
```
>>> assert common_attribute == 'Name'
```


Tests for q5.4 (points: 2)
```
>>> assert average_cost == 6.89
```


Tests for q5.5 (points: 2)
```
>>> assert top_rated_burritos.shape == (14, 7)
```


Tests for q5.6 (points: 2)
```
>>> assert average_cost_top_rated == 7.51
```


Tests for q5.7 (points: 3)
```
>>> assert top_rated_Califoria_count == 12
```
