```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
```

## <u>Name:</u> [Your name] and (your onyen)

Names and Onyens of fellow students you discussed Homework3 problems and ideas:

- Students: 

We encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 3: Pandas, Plots, and Statistics  (100 points)

Please complete this notebook by filling in the cells provided. Be sure to run the cell above that loads the autograder, and be sure to save your notebook before generating your submission by running the last cell.

!! Please submit a PDF of your notebook and upload it to Gradescope along with the generated zip file.  If no PDF submitted, **points will be deducted**.  You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop.  You can then drag and drop the zip & pdf file together to the upload box on Gradescope.

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!

Note that for homework, the inline autograder is useful feedback that you should pay attention to, but passing the autograder for a problem **DOES NOT GUARANTEE** you will get full credit on the problem, and you could still get 0 points for a problem that passes the inline autograder.  Also note that some problems will be manually graded.


## Problem 1. Pandas Lightning Round (15 points)
For each of the sentences below, write the corresponding Python code using Pandas.

Don't change any of the code that is provided.  Answers may require more than a single line of code.

NOTE!  As mentioned above, the inline autograder tests do NOT fully test your solution, and for MOST autograded problems additional tests are performed when submitting to Gradescope.  You could still get 0 points on a problem even if you pass the inline autograder test.  We highly recommend submitting to Gradescope early to check your solution.  There is no limit to the number of submissions.

**Problem 1.1 (1 point).** Load the file named `ROUSes.csv` into a table named `start_off_easy`.

```python
import numpy as np
import pandas as pd
rouses_df = pd.read_csv('ROUSes.csv')
```

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.1")
```

**Problem 1.2 (1 point).** Drop columns `'Length'` and `'Age'` from the table named `rouses_df` and assign the result to a table named `rouses_edited_df`.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.2")
```

**Problem 1.3 (1 point).** Assign to a variable named `max_age_first_10` the maximum `Age` among the first ten rows of the table `rouses_df`.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.3")
```

**Problem 1.4 (1 point).** To a variable named `lucky18`, assign the value of `Length` from the row with index `18` from the table `rouses_df`.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.4")
```

**Problem 1.5 (1 point).** To a variable named `average_weight_last15_rounded`, assign the value of the average of the `Weight` of the last fifteen rows of the table `rouses_df`, rounded to the nearest tenth. 

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.5")
```

**Problem 1.6 (2 points).** Assign to a variable named `younger_than_10` the number of ROUSes in table `rouses_df` that are younger than eleven.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.6")
```

**Problem 1.7 (2 points).** To a variable named `older_than_11_or_sleepy` assign the number of ROUSes in table `rouses_df` that are either older than eleven or are `Sleepy`.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.7")
```

**Problem 1.8 (2 points).** Assign to a variable named `silver_medal` the `Temperament` of the second-youngest ROUS in table `rouses_df`.  For ties in `Age`, use shortest `Length` to break the tie.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.8")
```

**Problem 1.9 (2 points).** Create a table that, for ROUSes in table `rouses_df` with the same temperament, has the medians of all the numeric columns. Assign it to a table named `medians_by_temperament`. 

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.9")
```

**Problem 1.10 (2 points).** Assign to a variable named `special_rouses`, how many ROUSes in table `rouses_df` are `Moody` and older than 8 and weigh less than 100.  

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q1.10")
```

## Problem 2: Burritos Revisited (22 points)

**Problem 2.1 (1 point)**. Reading and cleaning tables.

Let's revisit the `burritos` data from hw2. 

In the next cell, write three lines of code: 

1) read the file `ratings.csv` and save the resulting table as a variable called `ratings`,
2) read the file `burritos_types.csv` and save the resulting table as a variable called `burritos_types`, and
3) drop the column named `Unnamed: 0` from `burritos_types` and save that resulting table as a variable called `burritos_types` (that is, overwrite `burritos_types`).

If this sounds familiar, it's because we've done all of this in hw2. 

```python
# Write your code here!

```

**Problem 2.2 (5 points)**. 

Take a look at the following plot:

![burritos_avg](burritos_avg.png)

1) (1 point) Calculate the average burrito cost per restuarant and save that resulting table as a variable called `mean_cost_per_restaurant`. **Hint**: First decide whether you need `ratings` or `burritos_types`.

2) (1 point) First determine whether this plot is a line plot, scatter plot, bar chart, or a histogram. Then try to create a plot that looks like this using the appropriate seaborns function. No worries if you don't get the `cost` legend on the corner or if you get a different color.

3) (1 point) Now add the extra things - x axis label, y axis label, and title. I used fontsize 14 for the axes labels and 16 for the title.

4) (1 point) Determine what the units of the x and y axes should be, and make sure to replace the (units??) placeholders in the labels with the appropriate text.

5) (1 point) Without doing any calculations or coding, what would you say a typical burrito at a typical restaurant would cost in this town? In other words, if you walk into a random restaurant and buy a random burrito, how much do you expect to pay? Write a short sentence to justify your answer, using one (or more) of the words mean, median, mode, quartile, or percentile. Choose whichever is appropriate.


```python
# Write your code here!

```

[Write your answer here!]

**Problem 2.3 (3 points)**. 

Here's another plot:

![burritos_count](burritos_count.png)

1. (1 point) Calculate the number of burrito types sold in each restaurant and save that result as a variable called `num_burritos_per_restaurant`.

2. (1 point) Try to make this plot. I used the same fontsizes as before. Again, no need to worry if you get a different color or get a legend in the corner.

3. (1 point) Without doing any calculations or coding, how many types of burritos would you say a typical restaurant in this town sells? Write a short sentence to justify your answer, using one (or more) of the words mean, median, mode, quartile, or percentile. Choose whichever is appropriate.


```python
# Write your code here!


```

[Write your answer here!]

**Problem 2.4 (7 points)**.

Now we have a different type of plot for restaurant ratings.

![ratings.png](ratings.png)

1. (2 points) Try to recreate this plot including the axis labels and title. I used the same fontsizes as the earlier questions. No worries if you get different colors or small cosmetic differences. **Hint**: First decide what type of plot this is, and which seaborn function you need to work with. Then decide which table you need to work with.
2. (2 points) Without doing any calculations or coding, do you notice any patterns in the data, i.e. any relationships between ratings from different platforms? Which platform do you "trust" more (many valid answers) and does this plot change your opinion on whose ratings to trust more (it's okay if it doesn't)? Make sure to justify your answers. There are many valid answers.
3. (2 point) Do you notice anything odd about the Yelp plot? Is it *missing* anything, compared to the other two? Tell us what's missing in this Yelp plot, and why that would be the case. **Hint**: Try running the `describe` function on this table and see if you notice anything interesting that could lead to this observation.
4. (1 point) Calculate the 3 statistical dispersion measures (discussed in class) for Yelp and Google. Make sure to show your work.

```python
# Write your code here!

```

[Write your answer here!]

**Problem 2.5 (5 points)**.

You liked your previous plot so much that you decided to make it again for burritos.

![burritos_cost.png](burritos_cost.png)
1. (1 point) Make this plot, including the axis labels and title. It's okay if there are small cosmetic differences. Fontsizes are the same as earlier.


2. (2 point) Without doing any calculations or coding, what would you say a typical burrito would cost in this town? Is your answer different from Problem 1.2? **Is the distribution plotted here the same as the distribution plotted in Problem 1.2**? Make sure to justify your answers.

3. (1 point) Because of the outlier way up high, the plot is being squished. Can you find out which menu item from which restaurant is causing this problem?

4. (1 point) What is the average burrito cost **without** this outlier? How much less (rounded to the nearest cent) is this value compared to the average burrito cost **with** the outlier?

```python
# Write your code here. Add as many new cells (code or markdown) as you need.

```

[Write your answer here!]

**Problem 2.6 (1 point)**. Translating insights into action.

Surprise! You're actually a consultant hired by someone trying to break into the burrito industry in California. You are analyzing this dataset to gain some insight so that you can give good strategy advice to your client. What advice would you give to your client in terms of 1) how many burrito types they should sell and 2) how much they should charge for each burrito? Why? Make sure to justify your answer by referring to plots and calculations you've made in previous problems. There are many valid answers that would get full credit as long as they are well-justified.

[Write your answer here!]

## Problem 3: Basketball (18 points)

So I may or may not be contractually obligated to write at least one homework problem about basketball while at UNC.

Using data from this website (https://www.sports-reference.com/cbb/schools/north-carolina/men/), let's take a look at the UNC Men's Basketball history.

**Problem 3.1 (3 points)**. 
1. (1 point) Read the file `unc_all.csv` and save the resulting table as a variable called `unc`. Next, read the file `unc_coaches.csv` and save the resulting table as a variable named `unc_coaches`.
3. (2 points) Take a look at the `Season` column in `unc`. Our goal is to create a new column called `Year` that is derived from the `Season` column such that it takes the "beginning" year of the season. For example, I want "2023-24" in Season -> 2023 in Year, "2022-23" in Season -> 2022 in Year, and so on. Note that I am being intentional with the use of quotation marks here; I want all values in my new Year column to be of data type `int`.
 
**Hint**: You've seen examples in lecture where we took the first names and last names of the NBA players, and took the first word of the long scientific names in penguin species. This used the `.str.split(...).str.get(...)` notation. Find out what to replace `...` with, and what should come before the `.str`.

**Hint**: You can change the datatype of a column by using the `.astype(...)` notation. Same hint as above.

```python
# Write your code here! Add as many new cells as you need.


```

**Problem 3.2 (6 points)**.

Copy paste and run the following code, which should generate a line plot with two lines. Make sure you understand what's going on in the code. The best way to do that is by playing around and seeing how removing or changing code changes the plot.

```Python
fig, ax = plt.subplots(1)

sns.lineplot(data=unc, x='Year', y='W', label='W')
sns.lineplot(data=unc, x='Year', y='L', label='L')

ax.set_xlabel('Season', fontsize=14)
ax.set_ylabel('Games', fontsize=14)
ax.set_title("UNC Men's Basketball", fontsize=16)
ax.legend(fontsize=12)
```


1. (0.5 point) Based on this plot, do you think the total number of games played per season is increasing, decreasing, or staying the same over the years? Write a short sentence to justify your answer.
2. (0.5 point) Check your answer by creating a new column in `unc` called `Games` that counts the total number of games played per season. Then create a new lineplot only with this column.
3. (1 point) Coming back to the original plot, do you notice any "eras" or changepoints where something in the data shifted? How many did you count and in which seasons did this occur? Write a short sentence(s) to justify your answer. There are many valid answers. No calculation or coding needed.
4. (1 point) Here's how you can visually mark the eras or the changepoints on the plot directly. For example, if you think 1980 was a significant year when something changed, you can add this line to your plotting code above:
```Python
# all the other plotting code goes here
ax.axvline(x=1980, linestyle='-.', color='black')
# all the other plotting code goes here
```
Make another version of the original plot with these **vertical** line(s) that indicate the changepoints you've identified earlier. If you need to, play around with the figsize parameter in subplots or any of the optional parameters for axvline (https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.axvline.html) to make the plot look pretty.

5. (0.5 point) Now make a new lineplot using the `W-L%` column in `unc` table. Add the vertical line(s) to this plot as well.

6. (0.5 point) Do you observe the same shifts as before? Did anything change? Why (or why not) did the insight you got from this data change?

7. (0.5 point) Do the `W-L%` columns in `unc` and `unc-coaches` tables have the same distributions? Why or why not?

8. (0.5 point) Calculate the mean and median of the `W-L%` columns in both `unc` and `unc-coaches` tables.

9. (1 point) Add **horizontal** lines for mean and median to the line plot from Problem 3.2.5. This time make sure to have the legend as shown here. The code you need to draw a horizontal line is:

```Python
# all the other plotting code goes here
ax.axhline(x=..., linestyle=..., color=..., label=...)
# all the other plotting code goes here
```

![W-L.png](W-L.png)

```python
# Copy paste and run the code here. Add new cells (code and markdown) as you go.


```

[Write your answers here!]

**Problem 3.3 (4 points)**.

I made the following histogram of wins (with the `stat='density'` parameter):
![wins_density.png](wins_density.png)

The bins correspond to the coaches, which you can check from the `unc_coaches` table.

1. (1 point) What should the unit be on the y axis?
2. (1 point) Based on this histogram, which coach won the **most games**? Write a short sentence to justify your answer.

3. (1 point) Based on this histogram, which coach won the **most games per year**? Write a short sentence to justify your answer.
4. (1 point) Based on this histogram (and this histogram only), who would you vote as the most successful coach in the history of UNC Men's Basketball? Write a short sentence to justify your answer.

[Write your answers here!]

**Problem 3.4 (5 points)**.

Let's make some barplots! 

1. (1 point) Make a barplot of the number of **wins** for each coach. Please use the horizontal orientation so that we can read the coach names easily. **Hint**: First decide which table has that information.

2. (1 point) Create a new column for the table called `W/Yr` that keeps track of how many **wins per year** each coach has accomplished. Then make a new barplot with this column.

3. (1 point) Compare the two barplots. What are the main differences you notice?

4. (2 points) Now compare the set of two barplots to the histogram from earlier. Which information is covered in both sets of charts, and which information is covered in only one of them?

```python
# Write your code here! Add new cells as you need.

```

[Write your answers here!]

## Problem 4. Storytelling (10 points)

Something we've tried to emphasize in the data communication unit is that the same data can be used to tell different stories, sometimes intentionally and sometimes unintentionally. Let's investigate that further.

**Problem 4.1 (2 points).** Readings.

First, start by reading these two blog posts that summarize good and bad data visualization techniques.

https://medium.com/agoda-engineering/10-common-data-visualization-mistakes-and-how-to-avoid-them-e3896fe8e104

https://www.oldstreetsolutions.com/good-and-bad-data-visualization

What are two new things you learned from these readings? Use 4~8 sentences, and be as specific as you can. For example, "I learned that data visualization is important" will not get any credit.

If you have any trouble accessing these websites, please find the pdf files uploaded to the canvas website.

[Write your answer here!]

**Problem 4.2 (5 points)**.

1. (0.5 point) Find an example of a good data visualization from the internet. If you can, take a screenshot or copy the url to the image and embed it here (https://www.codecademy.com/resources/docs/markdown/images). If you can't, just adding a link here that gets you to the graphic is also fine. Although we didn't cover it in class, interactive visualizations are also welcome. If you're stuck, a good place to start might be newpaper websites (e.g. New York Times).
3. (0.5 point) Write down what the data behind the graphics is (e.g. US census, survey among some population) and what the source is. Your visualization (if it's any good) should mention where the data came from.
4. (1 point) Summarize in one sentence what you think this plot is trying to tell you. What do you think the creator of this visualization wanted you to know about the data, if you could take away just one thing?
5. (3 points) Write two things you like about this plot or why you thought this is an example of good data visualization. Write one thing you would like to change to improve the plot. This improvement could be for you specifically or in general.

[Write your answer here!]

**Problem 4.3 (3 points)**.

Pretend you are a nefarious villain named Puppycat who wants to mess with the world. 

You really don't like what the original plot was trying to tell you, i.e. what you wrote down in Problem 4.2.3. Therefore, you decide to create an intentionally *bad* and *misleading* data visualization from the *same set of data* (i.e. what you wrote down in Problem 4.2.2.) that tells a *completely different*, or even the *opposite* and *contradicting* message.

For example, let's say the original plot you picked shows that the unemployment rate in the US is decreasing, and therefore the US economy is doing well according to this measure. Evil Puppycat (that's you) wants to create a new plot that shows that the unemployment rate is increasing, or at the very least staying the same, using the same dataset.

Now, please explain how you whould achieve that in 3~6 sentences. Be as specific as you can by referring to specific parts of the plot. If you're stuck, referring back to the blog posts from Problem 4.1 might be helpful.

[Write your answer here!]

## Problem 5. Why Data Is Never Raw (10 points)

Read this essay by Nick Barrowman, published by the New Atlantis in 2018. 

https://www.thenewatlantis.com/publications/why-data-is-never-raw

Afterwards, write a short essay (100~300 words) on the following prompt.

What do you think this quote was trying to say about data? And how is this related to the data science life cycle model discussed in lecture (if any)?
    
    In the memorable words of Geoffrey Bowker, informatics professor at the University of California, Irvine, “Raw data is both an oxymoron and a bad idea; to the contrary, data should be cooked with care.”

Try to connect your thoughts to the rest of the article and what's been discussed in class so far. You are also allowed to (encouraged!) to bring up any of your own examples. This problem will be graded on relevance to the prompt, structure and clarity in writing. Vague, generic answers are not likely to get full credit.

[Write your answer here!]

## Problem 6. Modeling (15 points)
This problem deals with determining whether a proposed model is reasonable for describing one attribute as a function of another for ROUSes.

First, you will be evaluating which of 3 models would best describe how `Length` is a function of `Age`.

Start by running the code below.

```python
import seaborn as sns
import matplotlib.pyplot as plt
rouses_model_df = pd.read_csv("ROUSes.csv")
```

**Problem 6.1 (1 point).** Starting off by visualizing the data is always a good option.  Use seaborn `scatterplot()` to plot the $Length$ (y axis) vs. $Age$ (x axis).

```python
# write your code here, replacing ...
...
```

**Problem 6.2 (3 points).** Now add a new feature (column) to `rouses_model_df` for each of the 3 models below, and fill in the feature with data generated by the model.  Name the features `'Length Model 1'`, `'Length Model 2'`, and `'Length Model 3'`.  Note that the first model has a square root; for this you can use `**0.5`.  Hint: if you are unsure of mathematical operation ordering, parentheses are highly recommended!

  Length Model 1 = 1.08 * $\sqrt{Age}$

  Length Model 2 = 1.0 + 0.128 * $Age$

  Length Model 3 = 1.0 + 0.00229 * $Age^2$



```python
# write your code here, replacing ...
...
```

**Problem 6.3 (1 points).** There are many ways to compare data generated from a model to the actual data.  One of the easiest is to overlay plots of the actual and generated data.  One very simple way to do this is to put multiple scatterplots in the same cell, one after the other; seaborn is smart enough to automatically change the color for each of the plots.  Do this in the below; your plot should look like this: \
 ![](LengthModel.png)

```python
# write your code here, replacing ...
...
```

**Problem 6.4 (1 point).**  Just by looking at the plots, which of the models would you say is the best fit to the actual data?  

_Type your answer here, replacing this text._

**Problem 6.5 (2 points).** Summary statistics can also be used to compare data.  For the `Length` data and the data generated by each model, use pandas to help you compute the range to the nearst tenth.  Assign the 4 range values to the corresponding variables in the cells below.

```python
length_range = ...
length_range_model1 = ...
length_range_model2 = ...
length_range_model3 = ...
```

```python
grader.check("q6.5")
```

**Problem 6.6 (2 points).** Now do the same for median.  In other words, for the `Length` data and the data generated by each model, use pandas to help you compute the median to the nearst tenth.  Assign the 4 median values to the corresponding variables in the cells below.

```python
length_median = ...
length_median_model1 = ...
length_median_model2 = ...
length_median_model3 = ...
print(length_median, length_median_model1, length_median_model2, length_median_model3)

```

```python
grader.check("q6.6")
```

**Problem 6.7 (1 point).** Considering the values you computed, which summary statistic (range or median) would be more useful in deciding the model that best approximates the data?  Why? 

_Type your answer here, replacing this text._

**Problem 6.8 (1 point).** Another very useful visualization tool for comparing data is boxplot.  First, execute the code below:

```python
sns.boxplot(data=rouses_model_df)
```

```python
new_argument = ...
sns.boxplot(data=new_argument)
```

**Problem 6.9 (3 point).** Using any technique you wish, decide which of the models below is the best model for `Weight`.  Show some kind of work or visualization in the Python cell, and write what model matches best in the cell below that.

  Weight Model 1 = 51 * $\sqrt{Age}$

  Weight Model 2 = 6.9 * $Age$

  Weight Model 3 = 0.125 * $Age^2$



```python
# show your work here
...
```

_Type your answer here, replacing this text._

## Problem 7. Causality vs. correlation (10 points)
An important part of data science is understanding the relationships between data.  **Correlation** is the degree to which data change with one another.  **Causality** implies a direct relationship where changes in one thing directly results in changes to another thing.  Data can be correlated without any causality relationship; in fact, this occurs often when the data both depend on a third variable.  For example, number of shark attacks and ice cream sales are often correlated, but one does not cause the other; they both depend on weather conditions.  

**Problem 7.1 (2 points).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value? Why or why not?


```python
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [1,4,7,10,2,5,8,9,6,3,0]
sns.scatterplot(x=x_data, y=y_data)
```

_Type your answer here, replacing this text._

**Problem 7.2 (2 points).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value? Why or why not?


```python
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [10,13,14,16,18,20,21,24,26,27,30]
sns.scatterplot(x=x_data, y=y_data)
```

_Type your answer here, replacing this text._

**Problem 7.3 (2 points).** For problem 4.2, can you make a statement about causality?  In other words, can you say the behavior of ***y*** is caused by the behavior of ***x***?   Why or why not?

_Type your answer here, replacing this text._

**Problem 7.4 (4 points).** For each of the following cases, decide whether it is an example of **causality**, **correlation**, or **neither**.  For each answer, explain why in a single sentence.




* Case A. Roulette wheel spins and crap table dice rolls.

_Type your answer here, replacing this text._

* Case B.\
 ![](this_is_fine.jpg)

_Type your answer here, replacing this text._

* Case C. Golf club swing speed and golf ball initial flight speed.

_Type your answer here, replacing this text._

* Case D. Shoe size and reading ability.

_Type your answer here, replacing this text._

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework3

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1 (points: 1)
```
>>> assert list(start_off_easy.columns) == ['Age', 'Length', 'Weight', 'Temperament']
```


Tests for q1.10 (points: 2)
```
>>> assert type(special_rouses) == int
```


Tests for q1.2 (points: 1)
```
>>> assert rouses_edited_df.shape == (29, 2)
```


Tests for q1.3 (points: 1)
```
>>> assert (type(max_age_first_10) == float) | (type(max_age_first_10) == np.float64)
```


Tests for q1.4 (points: 1)
```
>>> assert (type(lucky18) == float) | (type(lucky18) == np.float64)
```


Tests for q1.5 (points: 1)
```
>>> assert (type(average_weight_last15_rounded) == float) | (type(average_weight_last15_rounded) == np.float64)
```


Tests for q1.6 (points: 2)
```
>>> assert type(younger_than_11) == int
```


Tests for q1.7 (points: 2)
```
>>> assert type(older_than_11_or_sleepy) == int
```


Tests for q1.8 (points: 2)
```
>>> assert type(silver_medal) == str
```


Tests for q1.9 (points: 2)
```
>>> assert medians_by_temperament.shape == (4, 3)
```


Tests for q6.5 (points: 2)
```
>>> assert length_range == 6.9
```


Tests for q6.6 (points: 2)
```
>>> assert length_median == 3.9
```
